```python
# Initialize Otter
!pip install -q otter-grader &> /dev/null
import otter
grader = otter.Notebook("python_quiz.ipynb")
```

Output:
```
/Users/handi7/anaconda3/lib/python3.7/site-packages/pandas/compat/_optional.py:138: UserWarning: Pandas requires version '2.7.0' or newer of 'numexpr' (version '2.6.8' currently installed).
  warnings.warn(msg, UserWarning)
```

In order for us to:

1. get a sense of how strong your python is
1. give you an example of how many of the assignments will work

Please give it an honest attempt. 

## **Question 1.** Extract Student Records from JSON

Read the provided JSON file "records.json" and extract a list of student records, with each record being a dictionary.

```python
!ls
```

Output:
```
README.md         python_quiz.ipynb records.json
```

```python
import json

file_path = "records.json"
with open(file_path, 'r') as f:
    data = json.load(f)

records = data['students']
print(records)
```

Output:
```
[{'name': '     Alice Johnson ', 'age': 20, 'courses': ['CS501', 'CS502', 'CS503']}, {'name': 'Bob    Smith ', 'age': 22, 'courses': ['CS505', 'CS504']}, {'name': '    Charlie Davis', 'age': 19, 'courses': ['CS501', 'CS508', 'CS506']}, {'name': '   Diana    Taylor', 'age': 21, 'courses': ['CS503', 'CS502', 'CS506', 'CS507']}, {'name': '   Bob   Dylan ', 'age': 21, 'courses': ['CS509', 'CS502', 'CS504']}]
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

## **Question 2.** Copy JSON Content Using shutil

Create a new JSON file called `clone-records.json` and copy the contents of 'records.json' into it using the shutil library.

```python
import shutil

# The grader is cloned-records.json instead of clone-records.json
# Typo here
dst = "cloned-records.json"
shutil.copyfile(src=file_path, dst=dst)
```

Output:
```
'cloned-records.json'
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

## **Question 3.** Sort Students by Age in Descending Order

Read the provided JSON file "records.json" and extract a list of student records and assign it to `sorted_records`, with each record being a dictionary. Use only one line of python.

```python
sorted_records = json.load(open(file_path))['students']

print(sorted_records)
```

Output:
```
[{'name': '     Alice Johnson ', 'age': 20, 'courses': ['CS501', 'CS502', 'CS503']}, {'name': 'Bob    Smith ', 'age': 22, 'courses': ['CS505', 'CS504']}, {'name': '    Charlie Davis', 'age': 19, 'courses': ['CS501', 'CS508', 'CS506']}, {'name': '   Diana    Taylor', 'age': 21, 'courses': ['CS503', 'CS502', 'CS506', 'CS507']}, {'name': '   Bob   Dylan ', 'age': 21, 'courses': ['CS509', 'CS502', 'CS504']}]
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

## **Question 4.** Clean and Format Student Names Using Regex

Write Python code to format the 'name' field in each dictionary from the list to the 'Firstname Lastname' format. Remove any extra white spaces between the first name and last name, as well as leading and trailing spaces in each name string. Store the result in `cleaned_records`. (Hint: Use Regex)

```python
import re

cleaned_records = []
for record in sorted_records:
    formatted_name = re.sub(r'\s+', '', record['name']).strip().title()
#     print(formatted_name)
    cleaned_records.append(formatted_name)
    
print(cleaned_records)
```

Output:
```
['Alicejohnson', 'Bobsmith', 'Charliedavis', 'Dianataylor', 'Bobdylan']
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

## **Question 5.** Add Course Count to Each Student Record

Create a new dataset called `courses_records` where each student record includes an attribute called `num_courses` representing the total number of courses the student is enrolled in. 

```python

print(courses_records)
```

```python
grader.check("q5")
```

## **Question 6.** Filter Students by Course Enrollment

Create a function that accepts a course name as input and returns a list of students who are enrolled in that course.

```python
def get_students(course_name):

print(get_students('CS506'), len(get_students('CS506')))
```

```python
grader.check("q6")
```

## **Question 7.** Calculate Average Age of Students

Calculate and return the average age of the students, rounding off to the nearest integer and assign it to a variable called `avg_age`.

```python
import numpy as np
def get_avg_age(student_list):

avg_age = ...
print(avg_age)
```

```python
grader.check("q7")
```

## **Question 8.** List of Unique Courses

Create a list of unique courses offered, as per the student records, and assign it to `total_courses`.

```python
import numpy as np
def get_unq_courses(student_list):

total_courses = ...
print(total_courses)
```

```python
grader.check("q8")
```

## **Question 9.** Group Students by Age 

Return a dictionary where each key is an age, and the value is a list of student names who are of that age, and assign it to a variable called `grouped_students`. Utilize itertools.groupby to group the list of students by their age.

```python
import itertools

def group_students_by_age(student_list):

grouped_students = ...
print(grouped_students)
```

```python
grader.check("q9")
```

## **Question 10.** Generate Course Enrollment Summary

Generate a dictionary where each key is a course name and the value is the number of students enrolled in that course.


```python
from collections import Counter

def get_course_dict(student_records):

course_enrollment = ...
print(course_enrollment)
```

```python
grader.check("q10")
```

Tests for q1 (points: 10)
```
>>> type(records) == list
True
```


Tests for q10 (points: 10)
```
>>> type(course_enrollment) == dict
True
```


Tests for q2 (points: 10)
```
>>> import os
>>> os.path.exists('cloned-records.json') and os.path.getsize('cloned-records.json') > 0
True
```


Tests for q3 (points: [0])
```
>>> len(sorted_records) == 5
True
```


Tests for q4 (points: 10)
```
>>> len(cleaned_records) == 5 and type(cleaned_records) == list
True
```


Tests for q5 (points: 10)
```
>>> len(courses_records) == 5 and type(courses_records) == list
True
```


Tests for q6 (points: 10)
```
>>> len(get_students('CS506')) == 2
True
```


Tests for q7 (points: 10)
```
>>> type(avg_age) == int
True
```


Tests for q8 (points: 10)
```
>>> type(total_courses) == np.ndarray
True
```


Tests for q9 (points: 10)
```
>>> type(grouped_students) == dict
True
```
